Uncertainty-modulated prediction errors in cortical microcircuits

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #2 (Public Review):

Summary:

This computational modeling study addresses the observation that variable observations are interpreted differently depending on how much uncertainty an agent expects from its environment. That is, the same mismatch between a stimulus and an expected stimulus would be less significant, and specifically would represent a smaller prediction error, in an environment with a high degree of variability than in one where observations have historically been similar to each other. The authors show that if two different classes of inhibitory interneurons, the PV and SST cells, (1) encode different aspects of a stimulus distribution and (2) act in different (divisive vs. subtractive) ways, and if (3) synaptic weights evolve in a way that causes the impact of certain inputs to balance the firing rates of the targets of those inputs, then pyramidal neurons in layer 2/3 of canonical cortical circuits can indeed encode uncertainty-modulated prediction errors. To achieve this result, SST neurons learn to represent the mean of a stimulus distribution and PV neurons its variance.

The impact of uncertainty on prediction errors in an understudied topic, and this study provides an intriguing and elegant new framework for how this impact could be achieved and what effects it could produce. The ideas here differ from past proposals about how neuronal firing represents uncertainty. The developed theory is accompanied by several predictions for future experimental testing, including the existence of different forms of coding by different subclasses of PV interneurons, which target different sets of SST interneurons (as well as pyramidal cells). The authors are able to point to some experimental observations that are at least consistent with their computational results. The simulations shown demonstrate that if we accept its assumptions, then the authors’ theory works very well: SSTs learn to represent the mean of a stimulus distribution, PVs learn to estimate its variance, firing rates of other model neurons scale as they should, and the level of uncertainty automatically tunes the learning rate, so that variable observations are less impactful in a high uncertainty setting.

Strengths:

The ideas in this work are novel and elegant, and they are instantiated in a progression of simulations that demonstrate the behavior of the circuit. The framework used by the authors is biologically plausible and matches some known biological data. The results attained, as well as the assumptions that go into the theory, provide several predictions for future experimental testing. The authors have taken into account earlier review comments to revise their paper in ways that enhance its clarity.

Weaknesses:

One weakness could be that the proposed theory does rely on a fairly large number of assumptions. However, there is at least some biological support for these. Importantly, the authors do lay out and discuss their key assumptions in the Discussion section, so readers can assess their validity and implications for themselves.

Thank you very much, we are very satisfied with this public review.

Reviewer #4 (Public Review):

Summary:

Wilmes and colleagues develop a model for the computation of uncertainty modulated prediction errors based on an experimentally inspired cortical circuit model for predictive processing. Predictive processing is a promising theory of cortical function. An essential aspect of the model is the idea of precision weighting of prediction errors. There is ample experimental evidence for prediction error responses in cortex. However, a central prediction of the theory is that these prediction error responses are regulated by the uncertainty of the input. Testing this idea experimentally has been difficult due to a lack of concrete models. This work provides one such model and makes experimentally testable predictions.

Strengths:

The model proposed is novel and well-implemented. It has sufficient biological accuracy to make useful and testable predictions.

Weaknesses:

One key idea the model hinges on is that stimulus uncertainty is encoded in the firing rate of parvalbumin positive interneurons. This assumption, however, is rather speculative and there is no direct evidence for this.

Thank you very much for this nice description. With regard to the weakness: it is true that the key idea hinges on uncertainty being encoded in the firing of inhibitory neurons. If it turns out that these inhibitory neurons are not PV neurons, however, the theory does not break down. The suggestion of PV neurons is fueled by the observation that PV neurons implement shunting and hence divisive inhibition and by the connectivity of PVs in the circuit. We discuss this in the discussion section: "To provide experimental predictions that are immediately testable, we suggested specific roles for SSTs and PVs, as they can subtractively and divisively modulate pyramidal cell activity, respectively. In principle, our theory more generally posits that any subtractive or divisive inhibition could implement the suggested computations. With the emerging data on inhibitory cell types, subtypes of SSTs and PVs or other cell types may turn out to play the proposed role."

Recommendations for the authors:

Reviewer #4 (Recommendations For The Authors):

(1) Line numbers would simplify reviewing.

We will add line numbers to our next submission.

(2) The existence of positive and negative PE was already suggested by Rao & Ballard.

We added the citation to the sentence "Because baseline firing rates are low in layer 2/3 pyramidal cells () positive and negative prediction errors were suggested to be represented by distinct neuronal populations [44,66],[...]" in the section "Computation of UPEs in cortical microcircuits".

(3) wekk should probably read well.

Indeed, thank you. We fixed it.

(4) Figure 4. legends A-C are mixed up. What are the two values of ¦s-u¦ in F and I - the same as in D and F.

Thank you, we fixed this.

(5) "representation neurons, the activity of which reflects the internal model". For consistency with the original definitions this should read "the activity of which reflects the internal representation". The internal "model" is the synaptic weights (or transformation between areas) - the activity of representation neurons (as the name implies) is the internal "representation".

Thank you, we changed it.

(6) "Mice trained in a predictable environment [...] [4]." This should read "reared" in an unpredictable environment, etc. Relatedly, the problem with this argument is that, the referenced paper argues that the mice never learned to predict and the reduced PE responses are a consequence of a reduction in prediction strength (these mice never - in life - had experience of visuomotor coupling). Better evidence might be the acute changes observed in normal mice (see e.g. Figure 3B in https://pubmed.ncbi.nlm.nih.gov/22681686/ However, another finding from the paper referenced is that in mice reared without visuomotor coupling, MM responses of SST interneurons are unchanged, while those in PV interneurons are completely absent. Would the authors model come to similar results if trained in an environment with (very) high uncertainty and then tested in a low uncertainty environment?

Thank you for pointing us to Figure 3B of Keller et al. 2012. We are now citing this result as it is indeed better evidence.

Thank you very much for your illuminating question and for pointing out that a mouse that never experienced a predictable visual flow may not have formed a model of the visual flow, and hence may not have any prediction about its visual experience. We haven’t considered this scenario in our paper before. So far, we only considered scenarios, in which it is possible to learn a prediction, i.e. to infer the mean from the sensory input. We now consider this other scenario in which the mouse that was reared in an unpredictable environment did not form a prediction and compare SST (1) and PV (2) activity in this mouse to one that learned to form a prediction, and added it to the section "Predictions for different cell types":

"Second, prediction error activity seems to decrease in less predictable, and hence more uncertain, contexts: in mice reared in a predictable environment [where locomotion and visual flow match, 42], error neuron responses to mismatches in locomotion and visual flow decreased with each day of experiencing these unpredictable mismatches. Third, the responses of SSTs and PVs to mismatches between locomotion and visual flow [4] are in line with our model (note that in this experiment the mismatches are negative prediction errors as visual flow was halted despite ongoing locomotion): In this study, SST responses decreased during mismatch, i.e. when the visual flow was halted, and there was no difference between mice reared in a predictable or unpredictable environment. In line with these observations, the authors concluded that SST responses reflected the actual visual input. In our model negative PE circuit, SSTs also reflect the actual stimulus input, which in our case was a whisker stimulus (SST rates in Fig. 6C and I reflect the stimuli (black and grey bar) in A and G, respectively) and SST rates are the same for high and low uncertainty (corresponding to mice reared in a predictable or unpredictable environment). In the same study, PV responses were absent towards mismatches in animals reared in an unpredictable environment [4]. The authors argued that mice reared in an unpredictable environment did not learn to form a prediction. In our model, the missing prediction corresponds to missing predictive input from the auditory domain (e.g. due to undeveloped synapses from the predictive auditory input). If we removed the predictive input in our model, PVs in the negative PE circuit would also be silent as they would not receive any of the excitatory predictive inputs."

(7) "Our model further posits the existence of two distinct subtypes of SSTs in positive and negative error circuits." There is some evidence for this: Figure 5a in https://pubmed.ncbi.nlm.nih.gov/36747710/

Thank you, we added this citation to the corresponding section.

Changes from the earlier version to the revised version

Abstract

Abstract

Understanding the variability of the environment is essential to function in everyday life. The brain must hence take uncertainty into account when updating its internal model of the world. The basis for updating the model are prediction errors that arise from a difference between the current model and new sensory experiences. Although prediction error neurons have been identified in layer 2/3 of diverse brain areas, how uncertainty modulates these errors and hence learning is, however, unclear. Here, we use a normative approach to derive how uncertainty should modulate prediction errors and postulate that layer 2/3 neurons represent uncertainty-modulated prediction errors (UPE). We further hypothesise that the layer 2/3 circuit calculates the UPE through the subtractive and divisive inhibition by different inhibitory cell types. By implementing the calculation of UPEs in a microcircuit model, we show that different cell types can compute the means and variances of the stimulus distribution. With local activity-dependent plasticity rules, these computations can be learned context-dependently, and allow the prediction of upcoming stimuli and their distribution. Finally, the mechanism enables an organism to optimise its learning strategy via adaptive learning rates.

Introduction

Earlier version: Decades of cognitive research indicate that our brain maintains a model of the world, based on which it can make predictions about upcoming stimuli [49, 11]. Predicting the sensory experience is useful for both perception and learning: Perception becomes more tolerant to uncertainty and noise when sensory information and predictions are integrated [59]. Learning can happen when predictions are compared to sensory information, as the resulting prediction error indicates how to improve the internal model. In both cases, the uncertainties (associated with both the sensory information and the internal model) should determine how much weight we give to the sensory information relative to the predictions, according to theoretical accounts. Behavioural and electrophysiological studies indicate that humans indeed estimate uncertainty and adjust their behaviour accordingly [59, 76, 28, 10, 47]. The neural mechanisms underlying uncertainty and prediction error computation are, however, less well understood. Recently, the activity of individual neurons of layer 2/3 cortical circuits in diverse cortical areas of mouse brains has been linked to prediction errors (visual, [43, 82, 22, 4, 26], auditory [17, 44], somatosensory [5], and posterior parietal [66]). Importantly, prediction errors could be associated with learning [41]. Prediction error neurons are embedded in neural circuits that consist of heterogeneous cell types, most of which are inhibitory. It has been suggested that prediction error activity results from an imbalance of excitatory and inhibitory inputs [34, 32], and that the prediction is subtracted from the sensory input [see e.g. 67, 4], possibly mediated by so-called somatostatin-positive interneurons (SSTs) [4]. How uncertainty is influencing these computations has not yet been investigated. Prediction error neurons receive inputs from a diversity of inhibitory cell types, the role of which is not completely understood. Here, we hypothesise that one role of inhibition is to modulate the prediction error neuron activity by uncertainty.

Revised version: Decades of cognitive research indicate that our brain maintains a model of the world, based on which it can make pre-dictions about upcoming stimuli [48, 11]. Predicting the sensory experience is useful for both perception and learning: Perception becomes more tolerant to uncertainty and noise when sensory information and predictions are integrated [58]. Learning can happen when predictions are compared to sensory information, as the resulting prediction error indicates how to improve the internal model. In both cases, the uncertainties (associated with both the sensory information and the internal model) should determine how much weight we give to the sensory information relative to the predictions, according to theoretical accounts. Behavioural and electrophysiological studies indicate that humans indeed estimate uncertainty and adjust their behaviour accordingly [58, 76, 28, 10, 46]. The neural mechanisms underlying uncertainty and prediction error computation are, however, less well understood. Recently, the activity of individual neurons of layer 2/3 cortical circuits in diverse cortical areas of mouse brains has been linked to prediction errors (visual, [42, 82, 22, 4, 26], auditory [17, 43], somatosensory [5], and posterior parietal [65]). Importantly, prediction errors could be associated with learning [40]. Prediction error neurons are embedded in neural circuits that consist of heterogeneous cell types, most of which are inhibitory. It has been suggested that prediction error activity results from an imbalance of excitatory and inhibitory inputs [34, 32], and that the prediction is subtracted from the sensory input [see e.g. 66, 4], possibly mediated by so-called somatostatin-positive interneurons (SSTs) [4]. How uncertainty is influencing these computations has not yet been investigated. Prediction error neurons receive inputs from a diversity of inhibitory cell types, the role of which is not completely understood. Here, we hypothesise that one role of inhibition is to modulate the prediction error neuron activity by uncertainty.

In this study, we use both analytical calculations and numerical simulations of rate-based circuit models with different inhibitory cell types to study circuit mechanisms leading to uncertainty-modulated prediction errors. First, we derive that uncertainty should divisively modulate prediction error activity and introduce uncertainty-modulated prediction errors (UPEs). We hypothesise that, first, layer 2/3 prediction error neurons reflect such UPEs. Second, we hypothesise that different inhibitory cell types are involved in calculating the difference between predictions and stimuli, and in the uncertainty modulation, respectively. Based on experimental findings, we suggest that SSTs and PVs play the respective roles. We then derive biologically plausible plasticity rules that enable those cell types to learn the means and variances from their inputs. Notably, because the information about the stimulus distribution is stored in the connectivity, single inhibitory cells encode the means and variances of their inputs in a context-dependent manner. Layer 2/3 pyramidal cells in this model hence encode uncertainty-modulated prediction errors context-dependently. We show that error neurons can additionally implement out-of-distribution detection by amplifying large errors and reducing small errors with a nonlinear fI-curve (activation function). Finally, we demonstrate that UPEs effectively mediate an adjustable learning rate, which allows fast learning in high-certainty contexts and reduces the learning rate, thus suppressing fluctuations in uncertain contexts.

Results

Normative theories suggest uncertainty-modulated prediction errors (UPEs)

In a complex, uncertain, and hence partly unpredictable world, it is impossible to avoid prediction errors. Some pre-diction errors will be the result of this variability or noise, other prediction errors will be the result of a change in the environment or new information. Ideally, only the latter should be used for learning, i.e., updating the current model of the world. The challenge our brain faces is to learn from prediction errors that result from new information, and less from prediction errors that result from noise. Hence, intuitively, if we learned that a kind of stimulus or context is very variable (high uncertainty), then a prediction error should have only little influence on our model. Consider a situation in which a person waits for a bus to arrive. If they learned that the bus is not reliable, another late arrival of the bus does not surprise them and does not change their model of the bus (Fig. 1A). If, on the contrary, they learned that the kind of stimulus or context is not very variable (low uncertainty), a prediction error should have a larger impact on their model. For example, if they learned that buses are reliable, they will notice that the bus is late and may use this information to update their model of the bus (Fig. 1A). This intuition of modulating prediction errors by the uncertainty associated with the stimulus or context is supported by both behavioural studies and normative theories of learning. Here we take the view that uncertainty is computed and represented on each level of the cortical hierarchy, from early sensory areas to higher level brain areas, as opposed to a task-specific uncertainty estimate at the level of decision-making in higher level brain areas (Fig. 1B) [see this review for a comparison of these two accounts: 84].

Distributed uncertainty-modulated prediction error computation in cortical circuits

A: A person who learned that buses are unreliable has a prior expectation, which can be described by a wide Gaussian distribution of expected bus arrival times. When the bus does not arrive at the scheduled time, this person is not surprised and remains calm, as everything happens according to their model of the world. On the other hand, a person who learned that buses are punctual, which can be described by a narrow distribution of arrival times, may notice that the bus is late and get nervous, as they expected the bus to be punctual. This person can learn from this experience. If they always took this particular bus, and their uncertainty estimate is accurate, the prediction error could indicate that the bus schedule changed. B: Models of uncertainty representation in cortex. Some models suggest that uncertainty is only represented in higher-level areas concerned with decision-making (left). In contrast, we propose that uncertainty is represented at each level of the cortical hierarchy (right, shown is the visual hierarchy as an example). C: a mouse learns the association between a sound (a) and a whisker deflection (s). The posterior parietal cortex (PPC) receives inputs from both somatosensory and auditory cortex. D: The whisker stimulus intensities are drawn from a Gaussian distribution with mean µ and standard deviation σ. E: Negative (left) and positive (right) prediction error circuit consisting of three cell types: layer 2/3 pyramidal cells (triangle), somatostatin-positive interneurons (SST, circle) and parvalbumin-positive interneurons (PV). SSTs represent the mean prediction in the postive circuit and the stimulus in the negative circuit, and PVs represent the variance.

Earlier version: Before we suggest how cortical circuits compute such uncertainty-modulated prediction errors, we consider the normative solution to a simple association that a mouse can learn. The setting we consider is to predict a somatosensory stimulus based on an auditory stimulus. The auditory stimulus a is fixed, and the subsequent somatosensory stimulus s is variable and sampled from a Gaussian distribution (s ∼ N (µ, σ), Fig. 1C,D). The optimal (maximum-likelihood) prediction is given by the mean of the stimulus distribution. Framed as an optimisation problem, the goal is to adapt the internal model of the mean such that the probability of observing samples s from the true distribution of whisker deflections is maximised given this model.

Revised version: Before we suggest how cortical circuits compute such uncertainty-modulated prediction errors, we consider the normative solution to a simple association that a mouse can learn. The setting we consider is to predict a somatosensory stimulus based on an auditory stimulus. The auditory stimulus a is fixed, and the subsequent somatosensory stimulus s is variable and sampled from a Gaussian distribution (s ∼ 𝒩 (µ, σ), Fig. 1C,D). The optimal (maximum-likelihood) prediction is given by the mean of the stimulus distribution. Framed as an optimisation problem, the goal is to adapt the internal model of the mean such that the probability of observing samples s from the true distribution of whisker deflections is maximised given this model.

Hence, stochastic gradient ascent learning on the log-likelihood suggests that with each observation s, the prediction (corresponding to the internal model of the mean) should be updated as follows to approach the maximum likelihood solution: where L is the likelihood.

According to this formulation, the update for the internal model should be the prediction error scaled inversely by the variance σ2. Therefore, we propose that prediction errors should be modulated by uncertainty.

Computation of UPEs in cortical microcircuits

Earlier version: How can cortical microcircuits achieve uncertainty modulation? Prediction errors can be positive or negative, but neuronal firing rates are always positive. Because baseline firing rates are low in layer 2/3 pyramidal cells [e.g., 58], positive and negative prediction errors were suggested to be represented by distinct neuronal populations [45], which is in line with experimental data [40]. We, therefore, decompose the UPE into a positive UPE+ and a negative UPE− component (Fig. 1C,D): where ⌊ … ⌋ + denotes rectification at 0.

Revised version: How can cortical microcircuits achieve uncertainty modulation? Prediction errors can be positive or negative, but neuronal firing rates are always positive. Because baseline firing rates are low in layer 2/3 pyramidal cells [e.g., 57], positive and negative prediction errors were suggested to be represented by distinct neuronal populations [44, 66], which is in line with experimental data [39]. We, therefore, decompose the UPE into a positive UPE+ and a negative UPE− component (Fig. 1C,D): where ⌊…⌋ + denotes rectification at 0.

Earlier version: It has been suggested that error neurons compute prediction errors by subtracting the prediction from the stimulus input (or vice versa) [4]. The stimulus input can come from local stimulus-encoding layer 2/3 cells [66]. Inhibitory interneurons provide the subtraction, resulting in an excitation-inhibition balance when they match [34]. To represent a UPE, error neurons need additionally be divisively modulated by the uncertainty. Depending on synaptic properties, such as reversal potentials, inhibitory neurons can have subtractive or divisive influences on their postsynaptic targets. Therefore, we propose that an inhibitory cell type that divisively modulates prediction error activity represents the uncertainty. We hypothesise, first, that in positive prediction error circuits, inhibitory interneurons with subtractive inhibitory effects represent the mean µ of the prediction. They probably either inherit the mean prediction or calculate it locally. Second, we hypothesise that inhibitory interneurons with divisive inhibitory effects represent the uncertainty σ2 of the prediction (Fig. 1C,D), which they calculate locally. A layer 2/3 pyramidal cell that receives these sources of inhibition then reflects the uncertainty-modulated prediction error (Fig. 1E,F).

Revised version: It has been suggested that error neurons compute prediction errors by subtracting the prediction from the stimulus input (or vice versa) [4]. The stimulus input can come from local stimulus-encoding layer 2/3 cells [65]. Inhibitory interneurons provide the subtraction, resulting in an excitation-inhibition balance when they match [34]. To represent a UPE, error neurons need additionally be divisively modulated by the uncertainty. Depending on synaptic properties, such as reversal potentials, inhibitory neurons can have subtractive or divisive influences on their postsynaptic targets. Therefore, we propose that an inhibitory cell type that divisively modulates prediction error activity represents the uncertainty. We hypothesise, first, that in positive prediction error circuits, inhibitory interneurons with subtractive inhibitory effects represent the mean µ of the prediction. They probably either inherit the mean prediction or calculate it locally. Second, we hypothesise that inhibitory interneurons with divisive inhibitory effects represent the uncertainty σ2 of the prediction (Fig. 1C,D), which they calculate locally. A layer 2/3 pyramidal cell that receives these sources of inhibition then reflects the uncertainty-modulated prediction error (Fig. 1E,F).

Earlier version: More specifically, we propose, first, that the SSTs are involved in the computation of the difference between predictions and stimuli, as suggested before [4]. This subtraction could happen on the apical dendrite. Second, we propose that the PVs provide the uncertainty modulation. In line with this, prediction error neurons in layer 2/3 receive subtractive inhibition from somatostatin (SST) and divisive inhibition from parvalbumin (PV) interneurons [81]. However, SSTs can also have divisive effects, and PVs can have subtractive effects, dependent on circuit and postsynaptic properties [71, 52, 15].

Revised version: More specifically, we propose, first, that the SSTs are involved in the computation of the difference between predictions and stimuli, as suggested before [4]. This subtraction could happen on the apical dendrite. Second, we propose that the PVs provide the uncertainty modulation. In line with this, prediction error neurons in layer 2/3 receive subtractive inhibition from somatostatin (SST) and divisive inhibition from parvalbumin (PV) interneurons [81]. However, SSTs can also have divisive effects, and PVs can have subtractive effects, dependent on circuit and postsynaptic properties [71, 51, 15].

In the following, we investigate circuits of prediction error neurons and different inhibitory cell types. We start by investigating in local positive and negative prediction error circuits whether the inhibitory cell types can locally learn to predict means and variances, before combining both subcircuits into a recurrent circuit consisting of both positive and negative prediction error neurons.

Local inhibitory cells learn to represent the mean and the variance given an associative cue

As discussed above, how much an individual sensory input contributes to updating the internal model should depend on the uncertainty associated with the sensory stimulus in its current context. Uncertainty estimation requires multiple stimulus samples. Therefore, our brain needs to have a context-dependent mechanism to estimate uncertainty from multiple past instances of the sensory input. Let us consider the simple example from above, in which a sound stimulus represents a context with a particular amount of uncertainty. In terms of neural activity, the context could be encoded in a higher-level representation of the sound. Here, we investigate whether the context representation can elicit activity in the PVs that reflects the expected uncertainty of the situation. To investigate whether the context provided by the sound can cause activity in SSTs and PVs that reflects the mean and the variance of the whisker stimulus distribution, respectively, we simulated a rate-based circuit model consisting of pyramidal cells and the relevant inhibitory cell types. This circuit receives both the sound and the whisker stimuli as inputs.

SSTs learn to estimate the mean

With our circuit model, we first investigate whether SSTs can learn to represent the mean of the stimulus distribution. In this model, SSTs receive whisker stimulus inputs s, drawn from Gaussian distributions (Fig. 2B), and an input from a higher level representation of the sound a (which is either on or off, see Eq. 9 in Methods). The connection weight from the sound representation to the SSTs is plastic according to a local activity-dependent plasticity rule (see Eq. 10). The aim of this rule is to minimise the difference between the activation of the SSTs caused by the sound input (which has to be learned) and the activation of the SSTs by the whisker stimulus (which nudges the SST activity in the right direction). The learning rule ensures that the auditory input itself causes SSTs to fire at the desired rate. After learning, the weight and the average SST firing rate reflect the mean of the presented whisker stimulus intensities (Fig. 2C-F).

SSTs learn to represent the mean context-dependently.

Illustration of the changes in the positive prediction error circuit. Thicker lines denote stronger weights. B: Two different tones (red, orange) are associated with two somatosensory stimulus distributions with different means (red: high, orange: low). C: SST firing rates (mean and std) during stimulus input. D: SST firing rates over time for low (orange) and high (red) stimulus means. E: Weights (mean and std) from sound a to SST for different values of µ. F: SST firing rates (mean and std) for different values of µ. Mean and std were computed over 1000 data points from timestep 9000 to 10000.

PVs learn to estimate the variance context-dependently

Earlier version: We next addressed whether PVs can estimate and learn the variance locally. To estimate the variance of the whisker deflections s, the PVs have to estimate σ2[s] = Es[(s − E [s])2] = Es[(s − µ)2]. To do so, they need to have access to both the whisker stimulus s and the mean µ. PVs in PPC respond to sensory inputs in diverse cortical areas [S1: 70] and are inhibited by SSTs in layer 2/3, which we assumed to represent the mean. Finally, for calculating the variance, these inputs need to be squared. PVs were shown to integrate their inputs supralinearly [12], which could help PVs to approximately estimate the variance.

Revised version: We next addressed whether PVs can estimate and learn the variance locally. To estimate the variance of the whisker deflections s, the PVs have to estimate σ2[s] = 𝔼s[(s − 𝔼[s])2] = 𝔼s[(s − µ)2]. To do so, they need to have access to both the whisker stimulus s and the mean µ. PVs in PPC respond to sensory inputs in diverse cortical areas [S1: 69] and are inhibited by SSTs in layer 2/3, which we assumed to represent the mean. Finally, for calculating the variance, these inputs need to be squared. PVs were shown to integrate their inputs supralinearly [12], which could help PVs to approximately estimate the variance.

Earlier version: In our circuit model, we next tested whether the PVs can learn to represent the variance of an upcoming whisker stimulus based on a context provided by an auditory input (Fig. 3A). Two different auditory inputs (Fig. 3B purple, green) are paired with two whisker stimulus distributions that differ in their variances (green: low, purple: high). PVs receive both the stimulus input as wekk as the inhibition from the SSTs, which subtracts the prediction of the mean (Eq. 11). The synaptic connection from the auditory input to the PVs is plastic according to the same local activity-dependent plasticity rule as the connection to the SSTs (Eq. 13). With this learning rule, the weight onto the PV becomes proportional to σ (Fig. 3C), such that the PV firing rate becomes proportional to σ2 on average (Fig. 3D). The average PV firing rate is exactly proportional to σ2 assuming a quadratic activation function φPV (x) (Fig. 3D-F,H) and monotonically increasing with σ2 with other choices of activation functions (Suppl. Fig. 10), both when the sound input is presented alone (Fig. 3D,E,H) or when paired with whisker stimulation (Fig. 3F). Notably, a single PV neuron is sufficient for encoding variances of different contexts because the context-dependent σ is stored in the connection weights.

Revised version: In our circuit model, we next tested whether the PVs can learn to represent the variance of an upcoming whisker stimulus based on a context provided by an auditory input (Fig. 3A). Two different auditory inputs (Fig. 3B purple, green) are paired with two whisker stimulus distributions that differ in their variances (green: low, purple: high). PVs receive both the stimulus input as well as the inhibition from the SSTs, which subtracts the prediction of the mean (Eq. 11). The synaptic connection from the auditory input to the PVs is plastic according to the same local activity-dependent plasticity rule as the connection to the SSTs (Eq. 13). With this learning rule, the weight onto the PV becomes proportional to σ (Fig. 3C), such that the PV firing rate becomes proportional to σ2 on average (Fig. 3D). The average PV firing rate is exactly proportional to σ2 assuming a quadratic activation function ϕP V (x) (Fig. 3D-F,H) and monotonically increasing with σ2 with other choices of activation functions (Suppl. Fig. 10), both when the sound input is presented alone (Fig. 3D,E,H) or when paired with whisker stimulation (Fig. 3F). Notably, a single PV neuron is sufficient for encoding variances of different contexts because the context-dependent σ is stored in the connection weights.

PVs learn to estimate the variance context-dependently.

A: Illustration of the changes in the positive prediction error circuit. Thicker lines denote stronger weights. B: Two different tones (purple, green) are associated with two somatosensory stimulus distributions with different variances (purple: high, green: low). C: Weights from sound a to PV over time for two different values of stimulus variance (high: σ = 0.8 (purple), low: σ = 0.4 (green)). D: PV firing rates over time given sound input (without whisker stimulus input) for low (green) and high (purple) stimulus variance. E: PV firing rates (mean and std) given sound input and whisker stimuli for low and high stimulus variance. F: PV firing rates (mean and std) during sound and stimulus input. G: Weights (mean and std) from sound a to PV for different values of σ. H: PV firing rates (mean and std) given sound input for different values of σ2. Mean and std were computed from 150000 data points from timestep 450000 to 600000.

To estimate the variance, the mean needs to be subtracted from the stimulus samples. A faithful mean subtraction is only ensured if the weights from the SSTs to the PVs (wPV,SST) match the weights from the stimuli s to the PVs (wPV,s). The weight wPV,SST can be learned to match the weight wPV,s with a local activity-dependent plasticity rule (see Suppl. Fig. 11 and Suppl. Methods).

The PVs can similarly estimate the uncertainty in negative prediction error circuits (Suppl. Fig. 12). In these circuits, SSTs represent the current sensory stimulus, and the mean prediction is an excitatory input to both negative prediction error neurons and PVs.

Calculation of the UPE in Layer 2/3 error neurons

Earlier version: Embedded in a circuit with subtractive and divisive interneuron types, layer 2/3 pyramidal cells could first compute the difference between the prediction and the stimulus in their dendrites, before their firing rate is divisively modulated by inhibition close to their soma. Layer 2/3 pyramidal cell dendrites can generate NMDA and calcium spikes, which cause a nonlinear integration of inputs. Hence, we took this into account and modelled the integration of dendritic activity as Idend = ⌊ wUPE,s s − wUPE,SST rSST ⌋ k with k determining the non-linearity. The total activity of prediction error neurons was modelled by The nonlinear integration of inputs is beneficial when the mean input changes and the current prediction differs strongly from the new mean of the stimulus distribution. For example, if the mean input increases strongly, the PV firing rate will increase for larger errors and inhibit error neurons more strongly than indicated by the learned variance estimate. The pyramidal nonlinearity compensates for this increased inhibition by PVs, such that in the end, layer 2/3 cell activity reflects an uncertainty-modulated prediction error (Fig. 4D-F) in both negative (Fig. 4A) and positive (Fig. 4B) prediction error circuits. A stronger nonlinearity (Fig. 4G-I) has the effect that error neurons elicit larger responses to larger prediction errors.

Revised version: Embedded in a circuit with subtractive and divisive interneuron types, layer 2/3 pyramidal cells could first compute the difference between the prediction and the stimulus in their dendrites, before their firing rate is divisively modulated by inhibition close to their soma. Layer 2/3 pyramidal cell dendrites can generate NMDA and calcium spikes, which cause a nonlinear integration of inputs. Hence, we took this into account and modelled the integration of dendritic activity as ⌊ Idend = wUPE,s s − wUPE,SST rSST⌋ k with k determining the non-linearity. The total activity of prediction error neurons was modelled by

Added in the revised version: The nonlinear integration of inputs is beneficial when the mean input changes and the current prediction differs strongly from the new mean of the stimulus distribution. For example, if the mean input increases strongly, the PV firing rate will increase for larger errors and inhibit error neurons more strongly than indicated by the learned variance estimate. The pyramidal nonlinearity compensates for this increased inhibition by PVs, such that in the end, layer 2/3 cell activity reflects an uncertainty-modulated prediction error (Fig. 4D-F) in both negative (Fig. 4A) and positive (Fig. 4B) prediction error circuits. A stronger nonlinearity (Fig. 4G-I) has the effect that error neurons elicit larger responses to larger prediction errors.

Calculation of the UPE in layer 2/3 error neurons

Earlier version: A: Illustration of the negative prediction error circuit. B: Illustration of the positive prediction error circuit. C: Illustration of error neuron with a nonlinear integration of inputs (k = 2). D: firing rate of the error neuron in the negative prediction error circuit (UPE−) as a function of σ for two values of |s − µ| after learning µ and σ. E: Rates of both UPE+ and UPE−-representing error neurons as a function of the difference between the stimulus and the mean (s − µ). F: firing rate of the error neuron in the positive prediction error circuit (UPE+) as a function of σ for two values of |s − µ| after learning µ and σ. G: Illustration of an error neuron with a non-linear activation function k = 2.5. H-J: same as D-F for error neurons with k = 2.5.

Revised version: A: Illustration of the negative prediction error circuit. B: Distributions with different standard deviations σ. C: Illustration of the positive prediction error circuit. D: Firing rate of the error neuron in the negative prediction error circuit (UPE−) as a function of σ for two values of |s − µ| after learning µ and σ. E: Rates of both UPE+ and UPE−-representing error neurons with a nonlinear activation function, where k = 2.0, as a function of the difference between the stimulus and the mean (s − µ). F: Firing rate of the error neuron in the positive prediction error circuit (UPE+) as a function of σ for two values of |s − µ | after learning µ and σ. G-I: same as D-F for error neurons with k = 2.5 with the same legend as in D-F.

To ensure a comparison between the stimulus and the prediction, the inhibition from the SSTs needs to match the excitation, which it is compared to, in the UPE neurons: In the positive PE circuit, the weights from the SSTs representing the prediction to the UPE neurons need to match the weights from the stimulus s to the UPE neurons. In the negative PE circuit, the weights from SSTs representing the stimulus to the negative UPE neurons need to match the weights from the mean representation to the UPE neurons, respectively. In line with previous proposals, error neuron activity signals the breaking of EI balance [34, 7]. With inhibitory plasticity (target-based, see Suppl. Methods), the weights from the SSTs can learn to match the incoming excitatory weights (Suppl. Fig. 13).

Interactions between representation neurons and error neurons

Earlier version: The theoretical framework of predictive processing includes both prediction error neurons and representation neurons, the activity of which reflects the internal model and should hence be compared to the sensory information. To make predictions for the activity of representation neurons, we expand our circuit model with this additional cell type. We first show that a representation neuron R can learn a representation of the stimulus mean given inputs from L2/3 error neurons. The representation neuron receives inputs from positive and negative prediction error neurons and from a higher level representation of the sound a (Fig. 5A). It sends its current mean estimate to the error circuits by either targeting the SSTs (in the positive circuit) or the pyramidal cells directly (in the negative circuit). Hence in this recurrent circuit, the SSTs inherit the mean representation instead of learning it. After learning, the weights wR,a from the sound representation to the representation neuron and the average firing rate of this representation neuron reflects the mean of the stimulus distribution (Fig.5B,C).

Revised version: The theoretical framework of predictive processing includes both prediction error neurons and representation neurons, the activity of which reflects the internal representation and should hence be compared to the sensory information. To make predictions for the activity of representation neurons, we expand our circuit model with this additional cell type. We first show that a representation neuron R can learn a representation of the stimulus mean given inputs from L2/3 error neurons. The representation neuron receives inputs from positive and negative prediction error neurons and from a higher level representation of the sound a (Fig. 5A). It sends its current mean estimate to the error circuits by either targeting the SSTs (in the positive circuit) or the pyramidal cells directly (in the negative circuit). Hence in this recurrent circuit, the SSTs inherit the mean representation instead of learning it. After learning, the weights wR,a from the sound representation to the representation neuron and the average firing rate of this representation neuron reflects the mean of the stimulus distribution (Fig.5B,C).

Learning the mean representation with UPEs

A: Illustration of the circuit. A representation neuron (turquoise) receives input from both positive and negative prediction error circuits (UPE+ and UPE−) and projects back to them. The UPE− has a negative impact on the firing rate of the representation neuron (rR). A weight wR,a from the higher level representation of the context given by sound a is learned. B: Weights wR,a over time for different values of µ ∈ (µ [1, 3, 5]). C: R firing rates given sound input for different values of µ (mean and std over 50000 data points from timestep 50000 to 100000, the end of the simulation). D: Activity of the different cell types (PV: light green, R: turquiose, UPE: black) and whisker stimulus samples (grey dots) over time. Learning the mean representation with PVs (light green) reflecting the MSE at the beginning, which is compensated by nonlinear activation of L2/3 neurons (black). The evolution of the mean rate of neuron R (turquoise) is similar to the perfect case in E. E: Same colour code as in D. Inset shows comparison to D. Learning the mean representation assuming PVs (light green) perfectly represent the variance.

As discussed earlier, pyramidal cells tend to integrate their dendritic inputs nonlinearly due to NMDA spikes. We here show that a circuit with prediction error neurons with a dendritic nonlinearity (as in Fig. 4) approximates an idealised circuit, in which the PV rate perfectly represents the variance (Fig. 5D,E, see inset for comparison of the two models). The dendritic nonlinearity can hence compensate for PV neuron dynamics. Also in this recurrent circuit, PVs learn to reflect the variance, as the weight from the sound representation a is learned to be proportional to σ (Suppl. Fig. 14).

Predictions for different cell types

Earlier version: Our model makes predictions for the activity of different cell types for positive and negative prediction errors (e.g. when a mouse receives whisker stimuli that are larger (Fig. 6A, black) or smaller (Fig. 6G, grey) than expected) in contexts associated with different amounts of uncertainty (e.g., the high-uncertainty (purple) versus the low-uncertainty (green) context are associated with different sounds). Our model suggests that there are two types of interneurons that provide subtractive inhibition to the prediction error neurons (presumably SST subtypes): in the positive prediction error circuit (SST+), they signal the expected value of the whisker stimulus intensity (Fig. 6B,H). in the negative prediction error circuit (SST−) they signal the whisker stimulus intensity (Fig. 6C,I). We further predict that interneurons that divisively modulate prediction error neuron activity represent the uncertainty (presumably PVs). Those do not differ in their activity between positive and negative circuits and may even be shared across the two circuits: in both positive and negative prediction error circuits, these cells signal the variance (Fig. 6D,J). L2/3 pyramidal cells that encode prediction errors signal uncertainty-modulated positive prediction errors (Fig. 6E) and uncertainty-modulated negative prediction errors (Fig. 6L), respectively. Finally, the existence of so-called internal representation neurons has been proposed [45]. In our case, those neurons represent the predicted mean of the associated whisker deflections. Our model predicts that upon presentation of an unexpected whisker stimulus, those internal representation neurons adjust their activity to represent the new whisker deflection depending on the variability of the associated whisker deflections: they adjust their activity more (given equal deviations from the mean) if the associated whisker deflections are less variable (see the next section and Fig. 7).

Revised version: Our model makes predictions for the activity of different cell types for positive and negative prediction errors (e.g. when a mouse receives whisker stimuli that are larger (Fig. 6A, black) or smaller (Fig. 6G, grey) than expected) in contexts associated with different amounts of uncertainty (e.g., the high-uncertainty (purple) versus the low-uncertainty (green) context are associated with different sounds). Our model suggests that there are two types of interneurons that pro-vide subtractive inhibition to the prediction error neurons (presumably SST subtypes): in the positive prediction error circuit (SST+), they signal the expected value of the whisker stimulus intensity (Fig. 6B,H). in the negative prediction error circuit (SST−) they signal the whisker stimulus intensity (Fig. 6C,I). We further predict that interneurons that divisively modulate prediction error neuron activity represent the uncertainty (presumably PVs). Those do not differ in their activity between positive and negative circuits and may even be shared across the two circuits: in both positive and negative prediction error circuits, these cells signal the variance (Fig. 6D,J). L2/3 pyramidal cells that encode prediction errors signal uncertainty-modulated positive prediction errors (Fig. 6E) and uncertainty-modulated negative prediction errors (Fig. 6L), respectively. Finally, the existence of so-called internal representation neurons has been proposed [44]. In our case, those neurons represent the predicted mean of the associated whisker deflections. Our model predicts that upon presentation of an unexpected whisker stimulus, those internal representation neurons adjust their activity to represent the new whisker deflection depending on the variability of the associated whisker deflections: they adjust their activity more (given equal deviations from the mean) if the associated whisker deflections are less variable (see the next section and Fig. 7).

Cell-type specific experimentally testable predictions

A: Illustration of the two experienced stimulus distributions with different variances that are associated with two different sounds (green, purple). The presented mismatch (MM) stimulus (black) is larger than expected (positive prediction error). B-F: Simulated firing rates of different cell types to positive prediction errors when a sound associated with high (purple) or low (green) uncertainty is presented. G: As in A. The presented mismatch stimulus (grey) is smaller than expected (negative prediction error). H-L: Firing rates of different cell types to the negative mismatch when a sound associated with high (purple) or low (green) uncertainty is presented. Because the firing rate predictions are equal for PV+ and PV−, we only show the results for PV+ in the figure.

Effective learning rate is automatically adjusted with UPEs

Earlier version: A,B: Firing rate over time of the representation neuron in a circuit with uncertainty-modulated prediction errors (gold) and in a circuit with unmodulated errors (black) in a low uncertainty setting (A) and a high uncertainty setting (B), C: standard deviation of the firing rate of the representation neuron in the low uncertainty setting (inset has a different scale, outer axis scale matches the one in D), D: standard deviation of the firing rate of the representation neuron in the high uncertainty setting. E: Standard deviation of the firing rate rR as a function of the standard deviation of the presented stimulus distribution σs. Standard deviations were computed over 100000 data points from timestep 100000 to 200000

Revised version: A,B: Firing rate over time of the representation neuron in a circuit with uncertainty-modulated prediction errors (gold) and in a circuit with unmodulated errors (black) in a low uncertainty setting and a high uncertainty setting (B), C: standard deviation of the firing rate of the representation neuron in the low uncertainty setting (inset has a different scale, outer axis scale matches the one in D), D: standard deviation of the firing rate of the representation neuron in the high uncertainty setting. E: Standard deviation of the firing rate rR as a function of the standard deviation of the presented stimulus distribution σs. Standard deviations were computed over 100000 data points from timestep 100000 to 200000

Earlier version: The following experimental results are compatible with our predictions: First, putative inhibitory neurons (narrow spiking units) in the macaque anterior cingulate cortex increased their firing rates in periods of high uncertainty [6]. These could correspond to the PVs in our model. Second, prediction error activity seems to be indeed lower for less predictable, and hence more uncertain, contexts: Mice trained in a predictable environment (where locomotion and visual flow match) were compared to mice trained in an unpredictable, uncertain environment [4, they saw a video of visual flow that was independent of their locomotion:]. Layer 2/3 activity towards mismatches in locomotion and visual flow was lower in the mice trained in the unpredictable environment.

Revised version: The following experimental results are compatible with our predictions: First, putative inhibitory neurons (narrow spiking units) in the macaque anterior cingulate cortex increased their firing rates in periods of high uncertainty [6]. These inhibitory neurons could correspond to the PVs in our model. Second, prediction error activity seems to decrease in less predictable, and hence more uncertain, contexts: in mice reared in a predictable environment [where locomotion and visual flow match 42], error neuron responses to mismatches in locomotion and visual flow decreased with each day of experiencing these unpredictable mismatches. Third, the responses of SSTs and PVs to mismatches between locomotion and visual flow [4] are in line with our model (note that in this experiment the mismatches are negative prediction errors as visual flow was halted despite ongoing locomotion): In this study, SST responses decreased during mismatch, i.e. when the visual flow was halted, and there was no difference between mice reared in a predictable or unpredictable environment. In line with these observations, the authors concluded that SST responses reflected the actual visual input. In our model negative PE circuit, SSTs also reflect the actual stimulus input, which in our case was a whisker stimulus (SST rates in Fig. 6C and I reflect the stimuli (black and grey bar) in A and G, respectively) and SST rates are the same for high and low uncertainty (corresponding to mice reared in a predictable or unpredictable environment). In the same study, PV responses were absent towards mismatches in animals reared in an unpredictable environment [4]. The authors argued that mice reared in an unpredictable environment did not learn to form a prediction. In our model, the missing prediction corresponds to missing predictive input from the auditory domain (e.g. due to undeveloped synapses from the predictive auditory input). If we removed the predictive input in our model, PVs in the negative PE circuit would also be silent as they would not receive any of the excitatory predictive inputs.

The effective learning rate is automatically adjusted with UPEs

To test whether UPEs can automatically adjust the effective learning rate of a downstream neural population, we looked at two contexts that differed in uncertainty and compared how the mean representation evolves with and without UPEs. Indeed, in a low-uncertainty setting, a new mean representation can be learned faster with UPEs (in comparison to unmodulated, Fig. 7A,C). In a high-uncertainty setting, the effective learning rate is smaller, and the mean representation is less variable than in the unmodulated case (Fig. 7B,D). The standard deviation of the firing rate increases only sublinearly with the standard deviation of the inputs (Fig. 7E). In summary, uncertainty-modulation of prediction errors enables an adaptive learning rate modulation.

UPEs ensure uncertainty-based weighting of prior and sensory information

Earlier version: Behavioural studies suggest that during perception humans integrate priors or predictions (p) and sensory information (s) in a Bayes-optimal manner [3, 61]. This entails that an internal neural representation (r), which determines perception, is achieved by weighting the two according to their uncertainties: where , is the uncertainty of the sensory information, and is the uncertainty of the prior.

Revised version: Behavioural studies suggest that during perception humans integrate priors or predictions (p) and sensory information (s) in a Bayes-optimal manner [3, 60]. This entails that an internal neural representation (r), which determines perception, is achieved by weighting the two according to their uncertainties: where , is the uncertainty of the sensory information, and is the uncertainty of the prior.

Earlier version: To obtain this weighting in the steady state, prediction errors from the lower area, the sensory prediction error (s − r), and from the local area, the representation prediction error (r − p), can be used to update the current representation [as in 67]. Maximising the log-likelihood (c.f. Eq. 22 in Methods and Eq. 36 in SI) yields an update of the representation by the difference between the bottom-up and top-down prediction errors. From this we obtain Eq. 4 by setting . Translating the update in Eq. 5 into our framework of UPE circuits reads as illustrated in Fig. 8A.

Revised version: To obtain this weighting in the steady state, prediction errors from the lower area, the sensory prediction error (s − r), and from the local area, the representation prediction error (r − p), can be used to update the current representation [as in 66]. Maximising the log-likelihood (c.f. Eq. 22 in Methods and Eq. 36 in SI) yields an update of the representation by the difference between the bottom-up and top-down prediction errors.

Added in the revised version: From this we obtain Eq. 4 by setting . Translating the update in Eq. 5 into our framework of UPE circuits reads as illustrated in Fig. 8A.

UPEs ensure uncertainty-based weighting of prior and sensory information.

A: Illustration of the hierarchical two-area model. A representation neuron (r) in area 𝓁 receives positive and negative UPEs from the area 𝓁 − 1 below (sensory prediction error, as in Fig. 5), and positive and negative UPEs from the same area (representation prediction error) with different signs, see Eq. 5. In this example, the uncertainty in area 𝓁 − 1 corresponds to the sensory uncertainty , and the uncertainty in the area 𝓁 above, corresponds to the prior uncertainty . Both uncertainties are represented by PV activity in the respective area. B,C: The simulation results from the hierarchical circuit model yield an uncertainty-based weighting of prior and sensory information as in Eq. 4. B: The activity r of the representation neuron as a function of the stimulus s for different amounts of prior uncertainty , when the sensory uncertainty is fixed . The histograms show the distribution of the sensory stimulus samples (grey) and the distribution of the activity of the representation neuron (green). C: The activity of the representation neuron as a function of the stimulus for different amounts of sensory uncertainty , when the prior uncertainty is fixed .

Earlier version: In our circuit, the representation neuron (green) receives the bottom-up errors, which are modulated by sensory uncertainty , as in the recurrent circuit model from the previous section. In addition, it receives errors between its current activity and a prior (a higher level expectation of the representation). The prior is set to the learned mean of the stimulus distribution and the errors are modulated by the learned prior uncertainty . We simulated this circuit, to which we presented stimuli drawn from Gaussian distributions as before. We varied the PV activity reflecting the sensory and prior uncertainty, and measured the activity of the representation neuron r to each stimulus. The higher the uncertainty of the prior information σs, the more the representation reflects the current stimulus (Fig. 8B, see also Eq. 4). The higher the uncertainty of the sensory information σs, the more the representation reflects the prior mean (Fig. 8C). This is a common behavioural effect, which is often referred to as central tendency effect or contraction bias [35, 37, 2, 57].

Revised version: In our circuit, the representation neuron (green) receives the bottom-up errors, which are modulated by sensory un-certainty , as in the recurrent circuit model from the previous section. In addition, it receives errors between its current activity and a prior (a higher level expectation of the representation). The prior is set to the learned mean of the stimulus distribution and the errors are modulated by the learned prior uncertainty . We simulated this circuit, to which we presented stimuli drawn from Gaussian distributions as before. We varied the PV activity reflecting the sensory and prior uncertainty, and measured the activity of the representation neuron r to each stimulus. The higher the uncertainty of the prior information σs, the more the representation reflects the current stimulus (Fig. 8B, see also Eq. 4). The higher the uncertainty of the sensory information σs, the more the representation reflects the prior mean (Fig. 8C). This is a common behavioural effect, which is often referred to as central tendency effect or contraction bias [35, 37, 2, 56].

To give a simple example how the prior uncertainty could come about in a dynamical environment, imagine noisy Gaussian sensory inputs with a fixed variance , and step changes in the mean, as in Fig. 7. On the lowest level 0, after faithful learning, the learned variance will represent the variance of the sensory input for a given mean. If in addition, the mean of the sensory input varies (as in Fig. 7), then on the level above will reflect how much the mean varies. The former corresponds to the sensory uncertainty, the latter to the environmental volatility, which increases the uncertainty about the prediction for the mean .

Discussion

Earlier version: Based on normative theories, we propose that the brain uses uncertainty-modulated prediction errors. In particular, we hypothesise that layer 2/3 prediction error neurons represent prediction errors that are inversely modulated by uncertainty. Here we showed that different inhibitory cell types in layer 2/3 cortical circuits can compute means and variances and thereby enable pyramidal cells to represent uncertainty-modulated prediction errors. We further showed that the cells in the circuit are able to learn to predict the means and variances of their inputs with local activity-dependent plasticity rules. Our study makes experimentally testable predictions for the activity of different cell types, PV and SST interneurons, in particular, prediction error neurons and representation neurons. Finally, we showed that circuits with uncertainty-modulated prediction errors enable adaptive learning rates, resulting in fast learning when uncertainty is low and slow learning to avoid detrimental fluctuations when uncertainty is high.

Revised version: Based on normative theories, we propose that the brain uses uncertainty-modulated prediction errors. In particular, we hypothesise that layer 2/3 prediction error neurons represent prediction errors that are inversely modulated by uncer-tainty. Here we showed that different inhibitory cell types in layer 2/3 cortical circuits can compute means and variances and thereby enable pyramidal cells to represent uncertainty-modulated prediction errors. We further showed that the cells in the circuit are able to learn to predict the means and variances of their inputs with local activity-dependent plasticity rules. Our study makes experimentally testable predictions for the activity of different cell types, PV and SST interneurons, in particular, prediction error neurons and representation neurons. Finally, we showed that circuits with uncertainty-modulated prediction errors enable adaptive learning rates, resulting in fast learning when uncertainty is low and slow learning to avoid detrimental fluctuations when uncertainty is high.

Earlier version: Our theory has the following notable implications: The first implication concerns the hierarchical organisation of the brain. At each level of the hierarchy, we find similar canonical circuit motifs that receive both feedforward (from a lower level) and feedback (from a higher level, predictive) inputs that need to be integrated. We propose that uncertainty is computed on each level of the hierarchy. This enables uncertainty estimates specific to the processing level of a particular area. Experimental evidence is so far insufficient to favour distributed uncertainty representation over the idea that uncertainty is only computed on the level of decisions in higher level brain areas such as the parietal cortex [46], orbitofrontal cortex [56], or prefrontal cortex [69]. Our study provides a concrete suggestion for a canonical circuit implementation and, therefore, experimentally testable predictions. The distributed account has clear computational advantages for task-flexibility, information integration, active sensing, and learning (see [84] for a recent review of the two accounts). Additionally, adding uncertainty-modulated prediction errors from different hierarchical levels according to the predictive coding model [67, 77] yields an uncertainty-based weighting of feedback and feedforward information, similar to a Bayes-optimal weighting, which can be reconciled with human behaviour [59]. Two further important implications result from storing uncertainty in the afferent connections to the PVs. First, this implies that the same PV cell can store different uncertainties depending on the context, which is encoded in the pre-synaptic activation. Second, fewer PVs than pyramidal cells are required for the mechanism, which is compatible with the 80/20 ratio of excitatory to inhibitory cells in the brain. The lower selectivity of interneurons in comparison to pyramidal cells could be a feature in prediction error circuits. Error neurons selective to similar stimuli are more likely to receive similar stimulus information, and hence similar predictions. Therefore, a circuit structure may have developed such that prediction error neurons with similar selectivity may receive inputs from the same inhibitory interneurons.

Revised version: Our theory has the following notable implications: The first implication concerns the hierarchical organisation of the brain. At each level of the hierarchy, we find similar canonical circuit motifs that receive both feedforward (from a lower level) and feedback (from a higher level, predictive) inputs that need to be integrated. We propose that uncertainty is computed on each level of the hierarchy. This enables uncertainty estimates specific to the processing level of a particular area. Experimental evidence is so far insufficient to favour distributed uncertainty representation over the idea that uncertainty is only computed on the level of decisions in higher level brain areas such as the parietal cortex [45], orbitofrontal cortex [55], or prefrontal cortex [68]. Our study provides a concrete suggestion for a canonical circuit implementation and, therefore, experimentally testable predictions. The distributed account has clear computational advantages for task-flexibility, information integration, active sensing, and learning (see [84] for a recent review of the two accounts). Additionally, adding uncertainty-modulated prediction errors from different hierarchical levels according to the predictive coding model [66, 77] yields an uncertainty-based weighting of feedback and feedforward information, similar to a Bayes-optimal weighting, which can be reconciled with human behaviour [58]. Two further important implications result from storing uncertainty in the afferent connections to the PVs. First, this implies that the same PV cell can store different uncertainties depending on the context, which is encoded in the pre-synaptic activation. Second, fewer PVs than pyramidal cells are required for the mechanism, which is compatible with the 80/20 ratio of excitatory to inhibitory cells in the brain. The lower selectivity of interneurons in comparison to pyramidal cells could be a feature in prediction error circuits. Error neurons selective to similar stimuli are more likely to receive similar stimulus information, and hence similar predictions. Therefore, a circuit structure may have developed such that prediction error neurons with similar selectivity may receive inputs from the same inhibitory interneurons.

Earlier version: We claim that the uncertainty represented by PVs in our theoretical framework corresponds to expected uncertainty that results from noise or irreducible uncertainty in the stimuli and should therefore decrease the learning rate. Another common source of uncertainty are changes in the environment, also referred to as the unexpected uncertainty. In volatile environments with high unexpected uncertainty, the learning rate should increase. We suggest that vasointestinal-peptide-positive interneurons (VIPs) could be responsible for signalling the unexpected uncertainty, as they respond to reward, punishment and surprise [64], which can be indicators of high unexpected uncertainty. They provide potent disinhibition of pyramidal cells [62], and also inhibit PVs in layer 2/3 [63]. Hence, they could increase error activity resulting in a larger learning signal. In general, interneurons are innervated by different kinds of neuromodulators [53, 65] and control pyramidal cell’s activity and plasticity [36, 25, 80, 79, 78]. Therefore, neuromodulators could have powerful control over error neuron activity and hence perception and learning.

Revised version: We claim that the uncertainty represented by PVs in our theoretical framework corresponds to expected uncertainty that results from noise or irreducible uncertainty in the stimuli and should therefore decrease the learning rate. Another common source of uncertainty are changes in the environment, also referred to as the unexpected uncertainty. In volatile environments with high unexpected uncertainty, the learning rate should increase. We suggest that vasointestinal-peptide-positive interneurons (VIPs) could be responsible for signalling the unexpected uncertainty, as they respond to reward, punishment and surprise [63], which can be indicators of high unexpected uncertainty. They provide potent disinhibition of pyramidal cells [61], and also inhibit PVs in layer 2/3 [62]. Hence, they could increase error activity resulting in a larger learning signal. In general, interneurons are innervated by different kinds of neuromodulators [52, 64] and control pyramidal cell’s activity and plasticity [36, 25, 80, 79, 78]. Therefore, neuromodulators could have powerful control over error neuron activity and hence perception and learning.

Earlier version: A diversity of proposals about the neural representation of uncertainty exist. For example, it has been suggested that uncertainty is represented in single neurons by the width [21], or amplitude of their responses [55], or implicitly via sampling [neural sampling hypothesis; 60, 9, 8], or rather than being represented by a single feature, can be decoded from the activity of an entire population [14]. While we suggest that PVs represent uncertainty to modulate prediction error responses, we do not claim that this is the sole representation of uncertainty in neuronal circuits.

Revised version: A diversity of proposals about the neural representation of uncertainty exist. For example, it has been suggested that uncertainty is represented in single neurons by the width [21], or amplitude of their responses [54], or implicitly via sampling [neural sampling hypothesis; [59, 9, 8], or rather than being represented by a single feature, can be decoded from the activity of an entire population [14]. While we suggest that PVs represent uncertainty to modulate prediction error responses, we do not claim that this is the sole representation of uncertainty in neuronal circuits.

Earlier version: Uncertainty estimation is relevant for Bayes-optimal integration of different sources of information, e.g., different modalities [multi-sensory integration; 18, 19] or priors and sensory information. Here, we present a circuit implementation for weighting sensory information versus priors according to their uncertainties by integrating uncertainty-modulated prediction errors. An alternative solution is to estimate uncertainty from the activity of prediction error neurons and use it to weight priors and sensory information [33], leading to contraction bias Hertäg and Clopath. [32] previously showed that the integration of prediction errors with sensory information in representation neurons can also lead to contraction bias, but without being dependent on uncertainty. Instead of modulating the error on each level by the uncertainty on that level as in our suggestion, one can also obtain a Bayes-optimal weighting by combining an unweighted top-down error with a bottom-error that is multiplicatively modulated by the prior uncertainty, and divisively modulated by the bottom-up uncertainty [29]. It has also been suggested that Bayes-optimal multi-sensory integration could be achieved in single neurons [19, 39]. Our proposal is complementary to these solutions in that uncertainty-modulated errors can be forwarded to other cortical and subcortical circuits at different levels of the hierarchy, where they can be used for inference and learning. It further allows for a context-dependent integration of sensory inputs.

Revised version: Uncertainty estimation is relevant for Bayes-optimal integration of different sources of information, e.g., different modal-ities [multi-sensory integration; 18, 19] or priors and sensory information. Here, we present a circuit implementation for weighting sensory information versus priors according to their uncertainties by integrating uncertainty-modulated prediction errors. An alternative solution is to estimate uncertainty from the activity of prediction error neurons and use it to weight priors and sensory information [33], leading to contraction bias Hertäg and Clopath. [32] previously showed that the integration of prediction errors with sensory information in representation neurons can also lead to contraction bias, but without being dependent on uncertainty. Instead of modulating the error on each level by the uncertainty on that level as in our suggestion, one can also obtain a Bayes-optimal weighting by combining an unweighted top-down error with a bottom-error that is multiplicatively modulated by the prior uncertainty, and divisively modulated by the bottom-up uncertainty [29]. It has also been suggested that Bayes-optimal multi-sensory integration could be achieved in single neurons [19, 38]. Our proposal is complementary to these solutions in that uncertainty-modulated errors can be forwarded to other cortical and subcortical circuits at different levels of the hierarchy, where they can be used for inference and learning. It further allows for a context-dependent integration of sensory inputs.

Earlier version: Multiple neurological disorders, such as autism spectrum disorder or schizophrenia, are associated with maladaptive contextual uncertainty-weighting of sensory and prior information [27, 51, 75, 50, 72]. These disorders are also associated with aberrant inhibition, e.g. ASD is associated with an excitation-inhibition imbalance [68] and reduced inhibition [31, 24]. Interestingly, PV cells, in particular chandelier PV cells, were shown to be reduced in number and synaptic strength in ASD [42]. Our theory provides one possible explanation of how deficits in uncertainty-weighting on the behavioural level could be linked to altered PVs on the circuit level.

Revised version: Multiple neurological disorders, such as autism spectrum disorder or schizophrenia, are associated with maladaptive contextual uncertainty-weighting of sensory and prior information [27, 50, 75, 49, 72]. These disorders are also associ-ated with aberrant inhibition, e.g. ASD is associated with an excitation-inhibition imbalance [67] and reduced inhibition [31, 24]. Interestingly, PV cells, in particular chandelier PV cells, were shown to be reduced in number and synaptic strength in ASD [41]. Our theory provides one possible explanation of how deficits in uncertainty-weighting on the behavioural level could be linked to altered PVs on the circuit level.

Earlier version: Finally, uncertainty-modulated errors could advance deep hierarchical neural networks. In addition to propagating gradients, propagating uncertainty may have advantages for learning. The additional information on uncertainty could enable calculating distances between distributions, which can provide an informative and parameter-independent metric for learning [e.g. natural gradient learning, 48].

Revised version: Finally, uncertainty-modulated errors could advance deep hierarchical neural networks. In addition to propagating gra-dients, propagating uncertainty may have advantages for learning. The additional information on uncertainty could enable calculating distances between distributions, which can provide an informative and parameter-independent metric for learning [e.g. natural gradient learning, 47].

Earlier version: To provide experimental predictions that are immediately testable, we suggested specific roles for SSTs and PVs, as they can subtractively and divisively modulate pyramidal cell activity, respectively. In principle, our theory more generally posits that any subtractive or divisive inhibition could implement the suggested computations. With the emerging data on inhibitory cell types, subtypes of SSTs and PVs or other cell types may turn out to play the proposed role. The model predicts that the divisive interneuron type, which we here suggest to be the PVs, receives a representation of the stimulus as an input. PVs could be pooling the inputs from stimulus-responsive layer 2/3 neurons to estimate uncertainty. The more the stimulus varies, the larger the variability of the pyramidal neuron responses and, hence, the variability of the PV activity. The broader sensory tuning of PVs [13] is in line with the model insofar as uncertainty modulation could be more general than the specific feature, which is more likely for low-level features processed in primary sensory cortices. PVs were shown to connect more to pyramidal cells with similar feature tuning [83]; this would be in line with the model, as uncertainty modulation should be feature-related. In our model, some SSTs deliver the prediction to the positive prediction error neurons. SSTs are already known to be involved in spatial prediction, as they underlie the effect of surround suppression [1], in which SSTs suppress the local activity dependent on a predictive surround.

Revised version: To provide experimental predictions that are immediately testable, we suggested specific roles for SSTs and PVs, as they can subtractively and divisively modulate pyramidal cell activity, respectively. In principle, our theory more generally posits that any subtractive or divisive inhibition could implement the suggested computations. With the emerging data on inhibitory cell types, subtypes of SSTs and PVs or other cell types may turn out to play the proposed role.

Added in the revised version: The model predicts that the divisive interneuron type, which we here suggest to be the PVs, receives a representation of the stimulus as an input. PVs could be pooling the inputs from stimulus-responsive layer 2/3 neurons to estimate uncertainty. The more the stimulus varies, the larger the variability of the pyramidal neuron responses and, hence, the variability of the PV activity. The broader sensory tuning of PVs [13] is in line with the model insofar as uncertainty modulation could be more general than the specific feature, which is more likely for low-level features processed in primary sensory cortices. PVs were shown to connect more to pyramidal cells with similar feature tuning [83]; this would be in line with the model, as uncertainty modulation should be feature-related. In our model, some SSTs deliver the prediction to the positive prediction error neurons. SSTs are already known to be involved in spatial prediction, as they underlie the effect of surround suppression [1], in which SSTs suppress the local activity dependent on a predictive surround.

Earlier version: In the model we propose, SSTs should be subtractive and PVs divisive. However, SSTs can also be divisive, and PVs subtractive dependent on circuit and postsynaptic properties [71, 52, 15]. This does not necessarily contradict our model, as circuits in which SSTs are divisive and PVs subtractive could implement a different function, as not all pyramidal cells are error neurons. Hence, our model suggests that error neurons which can calculate UPEs should have similar physiological properties to the layer 2/3 cells observed in the study by [81].

Revised version: In the model we propose, SSTs should be subtractive and PVs divisive. However, SSTs can also be divisive, and PVs subtractive dependent on circuit and postsynaptic properties [71, 51, 15]. This does not necessarily contradict our model, as circuits in which SSTs are divisive and PVs subtractive could implement a different function, as not all pyramidal cells are error neurons. Hence, our model suggests that error neurons which can calculate UPEs should have similar physiological properties to the layer 2/3 cells observed in the study by [81].

Earlier version: Our model further posits the existence of two distinct subtypes of SSTs in positive and negative error circuits. Indeed, there are many different subtypes of SSTs. SST is expressed by a large population of interneurons, which can be further subdivided. There is, for example, a type called SST44, which was shown to specifically respond when the animal corrects a movement [30]. Our proposal is hence aligned with the observation of functionally specialised subtypes of SSTs. Importantly, the comparison between stimulus and prediction needs to happen before the divisive modulation. Although our model does not make assumptions about the precise dendritic location of this comparison, we suggest this to happen on the apical dendrite, as top-down inputs and SST synapses arrive there. SSTs receive top-down inputs [54], which could provide the prediction to be subtracted in negative prediction error circuits.

Revised version: Our model further posits the existence of two distinct subtypes of SSTs in positive and negative error circuits. Indeed, there are many different subtypes of SSTs [70]. SST is expressed by a large population of interneurons, which can be further subdivided. There is, for example, a type called SST44, which was shown to specifically respond when the animal corrects a movement [30]. Our proposal is hence aligned with the observation of functionally specialised subtypes of SSTs. Importantly, the comparison between stimulus and prediction needs to happen before the divisive modulation. Although our model does not make assumptions about the precise dendritic location of this comparison, we suggest this to happen on the apical dendrite, as top-down inputs and SST synapses arrive there. SSTs receive top-down inputs [53], which could provide the prediction to be subtracted in negative prediction error circuits.

To compare predictions and stimuli in a subtractive manner, the encoded prediction/stimulus needs to be translated into a direct variable code. In this framework, we assume that this can be achieved by the weight matrix defining the synaptic connections from the neural populations representing predictions and stimuli (possibly in a population code).

To enable the comparison between predictions and sensory information via subtractive inhibition, we pointed out that the weights of those inputs on the postsynaptic neuron need to match. This essentially means that there needs to be a balance of excitatory and inhibitory inputs. Such an EI balance has been observed experimentally [73]. And it has previously been suggested that error responses are the result of breaking this EI balance [34, 7]. Heterosynaptic plasticity is a possible mechanism to achieve EI balance [20]. For example, spike pairing in pre- and postsynaptic neurons induces long-term potentiation at co-activated excitatory and inhibitory synapses with the degree of inhibitory potentiation depending on the evoked excitation [16], which can normalise EI balance [20].

Conclusion

To conclude, we proposed that prediction error activity in layer 2/3 circuits is modulated by uncertainty and that the diversity of cell types in these circuits achieves the appropriate scaling of the prediction error activity. The proposed model is compatible with Bayes-optimal behaviour and makes predictions for future experiments.

Methods

Derivation of the UPE

Added in the revised version: The goal is to learn to maximise the log-likelihood:

Earlier version: The goal is to learn to maximise the log-likelihood: We consider the log-likelihood for one sample s of the stimulus distribution: Stochastic gradient ascent on the log-likelihood gives the update for :

Revised version: We consider the log-likelihood for one sample s of the stimulus distribution:

Added in the revised version: Stochastic gradient ascent on the log-likelihood gives the update for :

Circuit model

Prediction error circuit

We modelled a circuit consisting of excitatory prediction error neurons in layer 2/3, and two inhibitory populations, corresponding to PV and SST interneurons.

Added in the revised version: Layer 2/3 pyramidal cells receive divisive inhibition from PVs [81]. We, hence, modelled the activity of prediction error neurons as where ϕ(x) is the activation function defined as:

Earlier version: Layer 2/3 pyramidal cells receive divisive inhibition from PVs [81]. We, hence, modelled the activity of prediction error neurons as where φ(x) is the activation function defined as: Idend = ⌊wUPE,s s − wUPE,SST rSST⌋k is the dendritic input current to the positive prediction error neuron (see section Neuronal dynamics below for rx and for the negative prediction error neuron, and Table 1 for wx). The nonlinearity in the dendrite is determined by the exponent k, which is by default k = 2, unless otherwise specified as in Fig. 4G-J. I0 > 1 is a constant ensuring that the divisive inhibition does not become excitatory, when σ < 1.0. All firing rates are rectified to ensure that they remain positive.

Revised version: Idend = ⌊wUPE,s s − wUPE,SST rSST⌋k is the dendritic input current to the positive prediction error neuron (see section Neuronal dynamics below for rx and for the negative prediction error neuron, and Table 1 for wx). The nonlinearity in the dendrite is determined by the exponent k, which is by default k = 2, unless otherwise specified as in Fig. 4G-J. I0 > 1 is a constant ensuring that the divisive inhibition does not become excitatory, when σ < 1.0. All firing rates are rectified to ensure that they remain positive.

Added in the revised version: In the positive prediction error circuit, in which the SSTs learn to represent the mean, the SST activity is determined by

Earlier version: In the positive prediction error circuit, in which the SSTs learn to represent the mean, the SST activity is determined by The SST activity is influenced (nudged with a factor β) by the somatosensory stimuli s, which provide targets for the desired SST activity. The connection weight from the sound representation to the SSTs wSST,a is plastic according to the following local activity-dependent plasticity rule [74]: where η is the learning rate, ra is the pre-synaptic firing rate, rSST is the post-synaptic firing rate of the SSTs, φ(x) is a rectified linear activation function of the SSTs.

Revised version: The SST activity is influenced (nudged with a factor β) by the somatosensory stimuli s, which provide targets for the desired SST activity. The connection weight from the sound representation to the SSTs wSST,a is plastic according to the following local activity-dependent plasticity rule [74]: where η is the learning rate, ra is the pre-synaptic firing rate, rSST is the post-synaptic firing rate of the SSTs, ϕ(x) is a rectified linear activation function of the SSTs.

The learning rule ensures that the auditory input alone causes SSTs to fire at their target activity. As in the original proposal [74], the terms in the learning rule can be mapped to local neuronal variables, which could be represented by dendritic (wSST,a ra) and somatic (rSST) activity.

Earlier version: The PV firing rate is determined by the input from the sound representation and the whisker stimuli, from which their mean is subtracted (, where the mean is given by ). The mean-subtracted whisker stimuli serve as a target for learning the weight from the sound representation to the PV . The PV firing rate evoles over time according to: where φPV(x) is a rectified quadratic activation function, defined as follows: The connection weight from the sound representation to the PVs wPV,a is plastic according to the same local activity-dependent plasticity rule as the SSTs [74]: The weight from the sound representation to the PV approaches σ (instead of µ as the weight to the SSTs), because the PV activity is a function of the mean-subtracted whisker stimuli (instead of the whisker stimuli as the SST activity), and for a Gaussian-distributed stimulus s ∼ N (s|µ, σ), it holds that E [⌊s − µ⌋+] ∝ σ.

Revised version: The PV firing rate is determined by the input from the sound representation (wPV+,a ra) and the whisker stimuli, from which their mean is subtracted (wPV+,s s − wPV+,SST+ rSST+, where the mean is given by rSST+). The mean-subtracted whisker stimuli serve as a target for learning the weight from the sound representation to the PV wPV+,a. The PV firing rate evoles over time according to: where ϕPV(x) is a rectified quadratic activation function, defined as follows:

Added in the revised version: The connection weight from the sound representation to the PVs wPV,a is plastic according to the same local activity-dependent plasticity rule as the SSTs [74]:

Added in the revised version: The weight from the sound representation to the PV wPV+,a approaches σ (instead of µ as the weight to the SSTs), because the PV activity is a function of the mean-subtracted whisker stimuli (instead of the whisker stimuli as the SST activity), and for a Gaussian-distributed stimulus s ∼ 𝒩 (s|µ, σ), it holds that 𝔼[⌊ s − µ⌋+] ∝ σ.

Recurrent circuit model

Earlier version: In the recurrent circuit, shown in Fig. 5, we added an internal representation neuron to the circuit with firing rate rR. In this circuit, the SSTs in the positive PE circuit inherit the mean representation from the representation neuron instead of learning it themselves, that is why they receive an input . The SSTs in the negative circuit inherit the stimulus representation and hence receive an input s In this recurrent circuit, the firing rate of each population ri where i ∈ [SST+, SST−, PV+, PV−, UPE+, UPE−, R] evolves over time according to the following neuronal dynamics. φ denotes a rectified linear activation function with saturation, φPV denotes a rectified quadratic activation function with saturation, defined in the section below. All firing rates are rectified to ensure that they remain positive.

Revised version: In the recurrent circuit, shown in Fig. 5, we added an internal representation neuron to the circuit with firing rate rR. In this circuit, the SSTs in the positive PE circuit inherit the mean representation from the representation neuron instead of learning it themselves, that is why they receive an input wSST+,RrR. The SSTs in the negative circuit inherit the stimulus representation and hence receive an input wSST−,s s In this recurrent circuit, the firing rate of each population ri where i ∈ [SST+, SST−, PV+, PV−, UPE+, UPE−, R] evolves over time according to the following neuronal dynamics. ϕ denotes a rectified linear activation function with saturation, ϕPV denotes a rectified quadratic activation function with saturation, defined in the section below. All firing rates are rectified to ensure that they remain positive.

Hierarchical predictive coding

Earlier version: The idea behind hierarchical predictive coding is that the brain infers or represents the causes of its sensory inputs using a hierarchical generative model [23]. Each level of the cortical hierarchy provides a prior for the mean of the lower level representation, with the top level representation rL being determined by the context. Noise enters in the sensory area by sampling a stimulus s = r0. In the sensory area, the variance is learned to match the variability of the external stimulus . The goal is to infer the set of latent representations r = (r1, .., rL−1), given the synaptic weights wf of the internal model and the top activity rL, that best reproduces the sensory inputs s.

Revised version: The idea behind hierarchical predictive coding is that the brain infers or represents the causes of its sensory inputs using a hierarchical generative model [23]. Each level of the cortical hierarchy provides a prior for the mean of the lower level representation, with the top level representation rL being determined by the context. Noise enters in the sensory area by sampling a stimulus s = r0. In the sensory area, the variance is learned to match the variability of the external stimulus . The goal is to infer the set of latent representations r = (r1, .., rL−1), given the synaptic weights w𝓁 of the internal model and the top activity rL, that best reproduces the sensory inputs s. To this end, we minimise the following energy: where ρ is a transfer function.

Removed in the revised version: To this end, we minimise the following energy: where ρ is a transfer function.

Earlier version: We obtain the update for rℳ with gradient descent on the energy with respect to rℳ: In our model, wℳ are scalars as they denote single weights.

Revised version: We obtain the update for r with gradient descent on the energy with respect to r𝓁 :

Added in the revised version: In our model, w𝓁 are scalars as they denote single weights.

Earlier version: We obtain the steady-state representation rℳ by setting its derivative in Eq. 21 to 0: We next consider, for simplicity, a threshold-linear transfer function ρ(wℳ−1rℳ) such that if wℳ−1rℳ < 0, then ρ(wℳ−1rℳ) = 0 and its derivative ρ′ = 0 or if wℳ−1rℳ ≥ 0 then ρ(wℳ−1rℳ) = wℳ−1rℳ and ρ′ = 1.

Revised version: We obtain the steady-state representation r by setting its derivative in Eq. 21 to 0:

Added in the revised version: We next consider, for simplicity, a threshold-linear transfer function ρ(w𝓁−1r𝓁) such that if w𝓁−1r𝓁 < 0, then ρ(w𝓁−1r𝓁) = 0 and its derivative ρ′ = 0 or if w𝓁 −1r𝓁 ≥ 0 then ρ(w𝓁 −1r𝓁) = w𝓁 −1r𝓁 and ρ′ = 1.

Earlier version: Solving Eq. 24 for rℳ and assuming rℳ ≥ 0, we get: See appendix for the general case with any transfer function and weight matrix.

Revised version: Solving Eq. 24 for r𝓁 and assuming r𝓁 ≥ 0, we get:

Added in the revised version: See appendix for the general case with any transfer function and weight matrix.

Hierarchical circuit model

Earlier version: In the hierarchical circuit model, the representation neuron does not only receive UPEs from the area below, but also from the current area. The UPEs from the area below are as defined in the recurrent circuit model, and the UPEs from the current area are defined accordingly as: The computations and parameters in each area are the same as for the recurrent circuit model above and in Fig. 5.

Revised version: In the hierarchical circuit model, the representation neuron does not only receive UPEs from the area below, but also from the current area.

Added in the revised version: The UPEs from the area below are as defined in the recurrent circuit model, and the UPEs from the current area are defined accordingly as:

Added in the revised version: The computations and parameters in each area are the same as for the recurrent circuit model above and in Fig. 5.

Synapses from the higher level representation of the sound a to R were plastic according to the following activity-dependent plasticity rules [74]. where ηPV = 0.01ηR.

Estimating the variance correctly

Earlier version: The PVs estimate the variance of the sensory input from the variance of the teaching input (s − µ), which nudges the membrane potential of the PVs with a nudging factor β. The nudging factor reduces the effective variance of the teaching input, such that in order to correctly estimate the variance, this reduction needs to be compensated by larger weights from the SSTs to the PVs (wPV,SST) and from the sensory input to the PVs (wPV,s). To determine how strong the weights ws = wPV,SST = wPV,s need to be to compensate for the downscaling of the input variance by β, we require that E [wa]2 = σ2 when the average weight change E [Δw] = 0. The learning rule for w is as follows: where and .

Revised version: The PVs estimate the variance of the sensory input from the variance of the teaching input (s − µ), which nudges the membrane potential of the PVs with a nudging factor β. The nudging factor reduces the effective variance of the teaching input, such that in order to correctly estimate the variance, this reduction needs to be compensated by larger weights from the SSTs to the PVs (wPV,SST) and from the sensory input to the PVs (wPV,s). To determine how strong the weights ws = wPV,SST = wPV,s need to be to compensate for the downscaling of the input variance by β, we require that 𝔼[wa]2 = σ2 when the average weight change 𝔼[Δw] = 0. The learning rule for w is as follows: where and .

Added in the revised version: Using that ϕ(u) = u2, the average weight change becomes:

Added in the revised version: Given our objective 𝔼[(wa)2] = σ2, we can write:

Added in the revised version: Then for 𝔼[Δw] = 0:

Earlier version: Using that φ(u) = u2, the average weight change becomes: Given our objective E [(wa)2] = σ2, we can write: Then for E [Δw] = 0: Here, we assumed that φ(u) = u2 instead of φ(u) = ⌊ u ⌋ 2. To test how well this approximation holds, we simulated the circuit for different values of β and hence ws, and plotted the PV firing rate rPV(a) given the sound input a and the weight from a to PV, wPV,a, for different values of β (Fig. 9). This analysis shows that the approximation holds for small β up to a value of β = 0.2.

Revised version: Here, we assumed that ϕ(u) = u2 instead of ϕ(u) = ⌊u⌋2. To test how well this approximation holds, we simulated the circuit for different values of β and hence ws, and plotted the PV firing rate rPV(a) given the sound input a and the weight from a to PV, wPV,a, for different values of β (Fig. 9). This analysis shows that the approximation holds for small β up to a value of β = 0.2.

For small β, and , the weight from a to PV approaches σ and the PV firing rate approaches σ2.

We initialised the circuit with the initial weight configuration in Tables 1 and 3 and neural firing rates were initialised to be 0 (ri(0) = 0 with i ∈ [SST+, SST−, PV+, PV−, UPE+, UPE−, R]). We then paired a constant tone input with N samples from the whisker stimulus distribution, the parameters of which we varied and are indicated in each Figure. Each whisker stimulus intensity was presented for D timesteps (see Table 2). All simulations were written in Python. Differential equations were numerically integrated with a time step of dt = 0.1.

Eliciting responses to mismatches (Fig. 4 and Fig. 6)

We first trained the circuit with 10000 stimulus samples to learn the variances in the a-to-PV weights. Then we presented different mismatch stimuli to calculate the error magnitude for each mismatch of magnitude s − µ.

Comparing the UPE circuit with an unmodulated circuit (Fig. 7)

Earlier version: To ensure a fair comparison, the unmodulated control has an effective learning rate that is the mean of the two effective learning rates in the uncertainty-modulated case.

Revised version: To ensure a fair comparison, the unmodu-lated control has an effective learning rate that is the mean of the two effective learning rates in the uncertainty-modulated case.